# Exercise 1

## Part 1 — First steps

### Step 1 — Make sure Jupyter Lab is running

If you have not already, install [Sunbeam](https://github.com/strawlab/sunbeam)
(see the main [README](../../../README.md)) and launch Jupyter Lab from it. You
should be reading this notebook inside that Jupyter Lab window.

### Step 2 — Blockly bird game

Play the games at https://blockly.games/.

Play the "Bird" game up to the highest level you can.

When you have solved a level, Blockly will show you a window saying
`Congratulations! You solved this level with N lines of JavaScript: XXXX Are
you ready for level Z?`

Copy the text you see into the code cell below, replacing the placeholder
text.

```python
# REPLACE THIS TEXT WITH THE JAVASCRIPT BLOCKLY SHOWS YOU

```

### Step 3 — Blockly with Python

We will also play with a version of Blockly for Python at
https://ulli.mingram.net/blockly_python.html

Use the turtle and draw an interesting shape like a triangle or a star. Copy
your Python program into the code cell below.

```python
# REPLACE THIS TEXT WITH YOUR PYTHON TURTLE PROGRAM

```

## Part 2 — Python basics: variables and functions

```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

## Problem 1 - variable assignment

Assign a value of 42 to the variable named `x`.

<!-- grade_id: cell-ae28d599cff14ed6, points: 0, grade: false, solution: true -->
```python
x=42
```

<!-- grade_id: cell-0e037ae0dbd9b1ad, points: 1, grade: true, solution: false -->
```python
# This checks that the above worked
assert ads_hash(x)=='73475cb40a'
```

## Problem 2 - solving exceptions

Modify the following so that it runs without an "exception" (an error in Python):


```python
y = 2 plus 2
```

<!-- grade_id: cell-c4b1b65a81fbe0ae, points: 0, grade: false, solution: true -->
```python
y=2+2
```

<!-- grade_id: cell-f25ca0617c9fefdf, points: 1, grade: true, solution: false -->
```python
# This checks that the above worked
assert ads_hash(y)=='4b227777d4'
```

## Problem 3 - using the Python Tutor

Run the following code at the [Python Tutor](http://pythontutor.com/).

1. Click on "Visualize your code and get live help now" in the Python Tutor website.
2. Copy this code into your clipboard:
```python
x = 42
y = x + 42
```
3. Paste it into the edit window of the Python Tutor.
4. Click the "Visualize Execution" button of the Python Tutor.
5. Click the "Next" button until all steps are done.

**For good insight, run all the problems in this exercise in the Python Tutor and carefully watch what is happening. This will help you understand what your programs are doing.**

How many variables are in the global frame? Assign the answer to the variable named `answer`.

<!-- grade_id: cell-c90583224070d298, points: 0, grade: false, solution: true -->
```python
answer=2
```

<!-- grade_id: cell-4a36c43c001f4d04, points: 1, grade: true, solution: false -->
```python
# This checks that the above worked
assert ads_hash(answer)=='d4735e3a26'
```

## Problem 4 - functions

Here is an example function named `double` that has one input named `x` and returns `x*2` (`x` times two).

```python
def double(x):
    return x*2
```

Remember that all functions in python are defined with the `def` word, then the name of the function, then zero, one, two, or more arguments in parentheses, then a colon. The body of the function must be in a block - a group of lines with the same indentation.

Now, make a function named `triple` that has one input named `x` and returns `x*3`.

<!-- grade_id: cell-69c4b1cf7dfd376f, points: 0, grade: false, solution: true -->
```python
def triple(x):
    return x*3
```

<!-- grade_id: cell-c2eb28444382a954, points: 1, grade: true, solution: false -->
```python
# This checks that the above worked
assert ads_hash(triple(3))=='19581e27de' 
assert ads_hash(triple(4))=='6b51d431df'
assert ads_hash(triple(5))=='e629fa6598'
```

## Problem 5 - functions with multiple arguments

Here is an example function named `my_func` that takes two inputs, `x` and `y` and does a bit of math on the inputs:

```python
def my_func(x,y):
    return x + 2*y
```

Make a function named `foo` which returns five times the first argument plus three times the second argument.

<!-- grade_id: cell-a7218905d3eb7d0d, points: 0, grade: false, solution: true -->
```python
# Write your answer here
def foo(a,b):
    return 5*a + 3*b
```

<!-- grade_id: cell-a417fcd5f3319f27, points: 1, grade: true, solution: false -->
```python
# This checks that the above worked
assert ads_hash(foo(3,10))=='811786ad1a'
assert ads_hash(foo(10,3))=='3e1e967e9b'
assert ads_hash(foo(13,2302))=='5f7635b6a1'
```

## Problem 6 - calling functions from functions

We can call functions from within functions.

Here are two example functions named `double` and `double_twice`

```python
def double(x):
    return x*2

def double_twice(x):
    y = double(x)
    z = double(y)
    return z
```

Re-use your function `triple` from above and make a function named `triple_thrice` which calls `triple` three times repeatedly on its input value.

<!-- grade_id: cell-aa0cfdb16d440f6b, points: 0, grade: false, solution: true -->
```python
def triple_thrice(x):
    return triple(triple(triple(x)))
```

<!-- grade_id: cell-ee28ca6d0fc146ee, points: 1, grade: true, solution: false -->
```python
assert ads_hash(triple_thrice(3))=='5316ca1c5d'
assert ads_hash(triple_thrice(4))=='9537f32ec7'
```

## Problem 7 - variable scopes

Functions are also useful because they let you "hide" variables so they do not "pollute" other parts of your code. In other words, functions let you make modular pieces of code which do not affect each other.

Consider this piece of code, and think about the value of `x` throughout the code.

```python
x = 43

def blah(x):
    x = x - 26
    return x

blah(2)

x
```

In the very last line, what is the value of `x`?

In the very last line, `x` has a value of `43`. However, inside the function there is a *local* variable, also named `x`, which can have a different value. Changes to this local variable `x` in the function do not affect the `x` in the *global frame*. This could also be called the *global scope*. "Frame" and "scope" mean similar things here.

Inside a function, variables from the outside frame can be *read* if there is no local variable in the function with the same name. Consider this code. Here we will read the variable `x` inside the function because we have no argument or other local variable named `x`.

```python
x = 43

def blah(y):
    z = x - y
    return z

q = blah(2)
```

What is the value of `q`?

The value of `q` here will be 41.

*Advanced*: Frames can be nested. For example, if you call a function named `function2` from another function named `function1`, the frame of `function2` will be created while the frame of `function1` continues to exist.

In this second example, during the call `blah(2)`, what value does the local variable `y` have inside `blah`? Assign your answer to the variable `y`.

<!-- grade_id: cell-cdc1a3f2940f028c, points: 0, grade: false, solution: true -->
```python
y=2
```

<!-- grade_id: cell-93d460ecce09dcfd, points: 1, grade: true, solution: false -->
```python
assert ads_hash(y)=='d4735e3a26'
```

## Problem 8 - return values of functions

To return a value, a function must use the keyword `return`. Otherwise, the special python type `None` is returned. Check these two functions out:

```python
def func_one(x):
    return x*10

def func_two(x):
    x*10
```

What is the result of `func_one(10)`? It is `100`. What is the result of `func_two(10)`? It is `None`, a special type in Python which means "nothing".

Taking this as a starting point, fix the `my_triple` function so that it returns the value of `x*3`.

```
def my_triple(x):
    x*3
```

<!-- grade_id: cell-3d76a38f83df69ce, points: 0, grade: false, solution: true -->
```python
def my_triple(x):
    return x*3
```

<!-- grade_id: cell-f70296c24b543ae5, points: 1, grade: true, solution: false -->
```python
assert ads_hash(my_triple(100))=='983bd614bb'
```

## Problem 9 - `print()` is not return

One thing that can be confusing is the difference between the return value of a cell in Jupyter notebook - which gets displayed automatically and the use of the built-in `print()` function.

```python
print(1234)
```

Will print `1234` but has no return value. This can be tricky in functions.

Consider this function

```python
def my_great_function(a):
    print(a*30)
```

What does this function return?

It returns `None`.

Now make a function called `my_even_better_function` which returns the value of its input times thirty.

<!-- grade_id: cell-25d97b85b08d5a2c, points: 0, grade: false, solution: true -->
```python
def my_even_better_function(a):
    return a*30
```

<!-- grade_id: cell-b2fc619e60513768, points: 1, grade: true, solution: false -->
```python
assert ads_hash(my_even_better_function(32))=='7f5642cd0c'
```

## Part 3 — Types and indexing

# Types

Please answer the questions about Python types. Here we will use the spelling that Python uses for the types. Thus, we write `int` (not e.g. `integer`) and `str` (not e.g. `string`).

For the special value `None`, Python's type is written as `type(None)`.

## Q1

What is the type of the following:

`10`

Assign your answer to the variable `my_type`.

<!-- grade_id: cell-4c5bb2e24cd0b4b8, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
my_type = int
```

<!-- grade_id: cell-a534d123d26a6e42, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(my_type)=='f918a1a5ca'
```

## Q2

What is the type of the following:
    
`10.0`

Assign your answer to the variable `my_type_2`.

<!-- grade_id: cell-c53f7736f3ef6a4d, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
my_type_2 = float
```

<!-- grade_id: cell-0997a02809ef51df, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(my_type_2)=='ec15c2bad4'
```

## Q3

What is the type of the following:
    
`"hello"`

Assign your answer to the variable `my_type_3`.

<!-- grade_id: cell-cecba2dc8f2a1fcb, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
my_type_3 = str
```

<!-- grade_id: cell-d8e48e520e48ecbb, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(my_type_3)=='17a0b2d655'
```

## Q4

What is the type of the following:
    
`None`

Assign your answer to the variable `my_type_4`.

<!-- grade_id: cell-01390d1edcadb200, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
my_type_4 = type(None)
```

<!-- grade_id: cell-81637082eb179284, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(my_type_4)=='b7253511ac'
```

## Q5

What is the type of the following:
    
`[1,2,3,4]`

Assign your answer to the variable `my_type_5`.

<!-- grade_id: cell-d6436534e3fbc839, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
my_type_5 = list
```

<!-- grade_id: cell-0e5bfe7fc5f7997b, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(my_type_5)=='981a391f15'
```

# Indexing

## Q1

Consider the following code. After this has run, what will the value of `x` be?

    y = [1,2,3]
    x = y[1]

Assign your answer to the variable `value_of_x`.

<!-- grade_id: cell-f5a2887e76c4321e, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
value_of_x = 2
```

<!-- grade_id: cell-5314eae929078d10, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(value_of_x)=='d4735e3a26'
```

## Q2

Consider the following code. After this has run, what will the value of `x` be?

    x = [1,2,3]
    x[2] = 42

Assign your answer to the variable `value_of_x`.    

<!-- grade_id: cell-8f0317f086f58de7, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
value_of_x = [1,2,42]
```

<!-- grade_id: cell-5aa9f92a6e6efaa0, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(value_of_x)=='ac009e066e'
```

## Q3

Consider the following code. After this has run, what will the value of `x` be?

    x = [1,2,3]
    y = x
    y[1] = 10

Assign your answer to the variable `value_of_x`.    

<!-- grade_id: cell-196bc244f1c8042b, points: 0, grade: false, solution: true -->
```python
# Type your answer here and then run this and the following cell.
value_of_x = [1,10,3]
```

<!-- grade_id: cell-39cc1a93f8204257, points: 1, grade: true, solution: false -->
```python
# If this runs without error, it means the answer in your previous cell was correct.
assert ads_hash(value_of_x)=='26796b1e40'
```