```python
# You must run this cell, but you can ignore its contents.

import hashlib

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

# Q1 SSH, command line, exit codes, etc 

SSH to your account at `python-course` (e.g. `ssh your-user-name@python-course.strawlab.org`) and follow the puzzle starting in the directory `/schnitzeljagd`.

You must be in the University of Freiburg network in order to be able to SSH correctly.

Your username is your the lowercase version of your lastname with some characters replaced according to the following python code:

```python
    username = username.replace('ä', 'ae')
    username = username.replace('ö', 'oe')
    username = username.replace('ü', 'ue')
    username = username.replace('ß', 'ss')
    username = username.replace('ó', 'o')
    username = username.replace(' ', '')
```

Your password is your matriculation number followed by "x" (e.g. `12345678x`).

Place the final answer in the variable `final_answer`.

<!-- grade_id: cell-d6b141313f4f3a09, points: 0, grade: false, solution: true -->
```python
final_answer = "lada0df89ksdjf8alkds9339akklafk93"
```

<!-- grade_id: cell-3d5b7fe600c4c6ba, points: 1, grade: true, solution: false -->
```python
assert ads_hash(final_answer) == 'dfb81c4fac'
```

# Q2 Starting `uv` environment from command line

0) Login to `python-course.strawlab.org` as described above.
1) Setup and activate the uv `course` environment as described in the lecture notebook (Section "Setting up our Python environment with `uv`").
2) Change to the directory `/shared-programs`.
3) Run the python program `program1.py`
4) Put the program console output into the variable `program1_output` below

<!-- grade_id: cell-6d9ed8460134bae3, points: 0, grade: false, solution: true -->
```python
program1_output=300
```

<!-- grade_id: cell-bc67d9f37b59741d, points: 1, grade: true, solution: false -->
```python
assert ads_hash(program1_output) == '983bd614bb'
```