# Exercise 4 — Coding with an AI assistant

The lecture was about working *effectively* with AI coding assistants: describing a
problem clearly, reading and running what you get back, and **verifying** it
instead of trusting it. This exercise makes you do exactly that.

You will use the university chat assistant at
[openwebui.uni-freiburg.de](https://openwebui.uni-freiburg.de). Use a "coding model" such as "UFR: Standard Coding". The central computing facility continually updates these generic names to use a specific coding model such as GPT-5.3, GLM-5.2, or Gemma 4. You are
**expected** to use the assistant here — but you are also expected to check every
answer it gives you. As the course policy says: you must be able to understand,
test, and explain any code you submit.

**How this exercise works.** As in previous exercises, for coding questions, replace the
`raise NotImplementedError()` line with your answer (which you may develop with
the assistant), then run that cell **and** the check cell that follows it. If a
check cell runs without an error, (or if you click the "Validate" button), your answer passed. A few cells ask you to
write a short reflection in words — these are manually reviewed.

**A note on the assistant.** The chat assistant only produces *text*. It cannot
see your files, run your code, or read your error messages unless you paste them
in. That is exactly why *you* are the one who has to run and verify everything.

When you are done, upload this `.ipynb` file to ILIAS. **Do not rename the
file.**

```python
# You must run this cell, but you can ignore its contents.
import hashlib
import matplotlib.pyplot as plt

def ads_hash(ty):
    """Return a unique string for input"""
    ty_str = str(ty).encode()
    m = hashlib.sha256()
    m.update(ty_str)
    return m.hexdigest()[:10]
```

## Part 1 — The workflow in miniature

The basic loop for working with an AI assistant is: **specify → generate → read →
run → test**. Let's practise it on something small.

### Q1

Open the chat assistant and ask it to write a Python function for you. Your task
is to get a working function called `gc_content` that computes the **GC content**
of a DNA sequence — the fraction of bases that are `G` or `C`.

Be specific in your prompt. The function must satisfy:

- it takes one argument, a string `seq` (for example `"GATTACA"`),
- it returns a `float`: the number of `G` and `C` bases divided by the total
  number of bases,
- it is **case-insensitive**, so `"g"` and `"G"` both count.

Paste the assistant's function into the cell below. Then **read it** — does it
actually do what was asked? — and run the check cell. If the check fails, that is
the assistant getting it wrong: tell it what went wrong and iterate.

<!-- grade_id: cell-ex04-gc-sol, points: 0, grade: false, solution: true -->
```python
def gc_content(seq: str) -> float:
    """
    Compute the GC content of a DNA sequence.

    The GC content is the fraction of bases that are G or C.
    The calculation is case-insensitive.

    Args:
        seq (str): A DNA sequence (e.g., "GATTACA").

    Returns:
        float: The fraction of G and C bases. Returns 0.0 if the sequence is empty.
    """
    if not seq:
        return 0.0
    # Convert to lowercase to handle case-insensitivity
    seq_lower = seq.lower()
    # Count G and C bases
    gc_count = seq_lower.count('g') + seq_lower.count('c')
    return gc_count / len(seq)
```

<!-- grade_id: cell-ex04-gc-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your gc_content function is correct.
assert ads_hash(round(gc_content("GGCC") * 1000)) == '4051017584'
assert ads_hash(round(gc_content("ATAT") * 1000)) == '5feceb66ff'
assert ads_hash(round(gc_content("AtGc") * 1000)) == '0604cd3138'   # case-insensitive
assert ads_hash(round(gc_content("GATTACA") * 1000)) == '00328ce57b'
```

## Part 2 — A real, messy dataset

Real data is rarely clean, and an assistant that cannot see your file will not
warn you about that. In this part you will load a real biological dataset and
compute a couple of numbers from it — carefully.

**The data.** `antarctic_penguins.csv` contains body measurements of penguins
from three species. It is a lightly reformatted version of the *Palmer Penguins*
dataset (Gorman, Williams & Fraser, 2014; packaged by Allison Horst). The file
sits in the same folder as this notebook. By the way, the Culmen is the ridge on the top of a bird's beak: https://en.wikipedia.org/wiki/Beak#Culmen.

Let's look at the first few raw lines before we do anything else — always a good
habit.

```python
fobj = open('antarctic_penguins.csv')
for line_num, line in enumerate(fobj.readlines()):
    line = line.strip()
    print(f"line {line_num}: '{line}'")
    if line_num > 4:
        break
```

### Q1 — Load the file into a dictionary

Notice, from the raw lines above, that this file is **not** a plain
comma-separated file. Ask the assistant to write code (using only the Python
standard library — no `pandas`) that reads `antarctic_penguins.csv` into a dict
called `data`, where each key is a column name and each value is the list of
entries in that column. The four measurement columns
(`culmen_length_mm`, `culmen_depth_mm`, `flipper_len_mm`, `body_mass_g (g)`)
should be lists of **numbers**; the other columns can stay as strings.

The assistant probably will *not* get this right on the first try, because it
cannot see the quirks of the file. When you run its code you will hit real
errors. **Read the error, paste it back to the assistant, and iterate** until the
code runs and the check passes. (Hint: look closely at how the fields are
separated, and at how the decimal point is written.)

<!-- grade_id: cell-ex04-load-sol, points: 0, grade: false, solution: true -->
```python
data = {}

with open('antarctic_penguins.csv') as fobj:
    lines = fobj.readlines()

# First line is the header, split on ';' to get column names
header = lines[0].strip().split(';')

# Columns that should be converted to numbers
numeric_cols = {'culmen_length_mm', 'culmen_depth_mm',
                'flipper_len_mm', 'body_mass_g (g)'}

# Initialize each column with an empty list
data = {col: [] for col in header}

# Process the remaining lines
for line in lines[1:]:
    line = line.strip()
    if not line:          # skip blank lines (e.g. trailing newline)
        continue
    values = line.split(';')
    for col, val in zip(header, values):
        if col in numeric_cols:
            # European decimal separator: comma -> dot
            val = val.replace(',', '.')
            # Use int when there's no decimal part, otherwise float
            data[col].append(int(val) if '.' not in val else float(val))
        else:
            data[col].append(val)
```

<!-- grade_id: cell-ex04-load-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your file was parsed correctly.
assert ads_hash(len(data)) == '7902699be4'                       # 7 columns
assert ads_hash(len(data['species'])) == '023849c389'            # 342 rows
assert ads_hash(round(data['culmen_length_mm'][0] * 10)) == 'a934c24475'  # first value 39.1
```

### Q2 — Average body mass (verify before you trust!)

Now compute the **mean body mass** of the penguins and store it in a variable
called `mean_body_mass`.

This looks trivial — and the assistant will happily give you a one-line answer
that *runs without any error*. But a number that runs is not the same as a number
that is correct. Real datasets often record missing measurements with special
"sentinel" codes rather than leaving them blank. **Inspect the body-mass values
yourself** (for example, print the smallest and largest values, or sort them)
before you trust your mean. Make sure `mean_body_mass` reflects only the genuine
measurements.

<!-- grade_id: cell-ex04-mean-sol, points: 0, grade: false, solution: true -->
```python
masses = [m for m in data['body_mass_g (g)'] if m != -999]
mean_body_mass = sum(masses) / len(masses)
```

<!-- grade_id: cell-ex04-mean-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your mean excludes the missing-data sentinel.
assert ads_hash(round(mean_body_mass)) == 'b68f7a167a'
```

### Q3 — Reflect (written answer, graded by a human)

In a sentence or two: the first mean you computed was almost certainly wrong.
What was actually in the data that made it wrong, how did you notice, and what did
you change? If the assistant helped you find or fix it, say how.

<!-- grade_id: cell-ex04-reflect1, points: 1, grade: true, solution: true -->
Excluded sentinel value of -999.

## Part 3 — Reading AI-generated code critically

An assistant produces *plausible-looking* code. Plausible is not the same as
correct. Below are three short functions of the kind an assistant might hand you.
Each one is **broken**. For each: read it, work out what is wrong (running it and
testing it against values you can check by hand is the fastest way), and write a
**corrected** version in the solution cell. You may ask the assistant for help —
but you must be able to explain the bug yourself.

### Bug A — a function that doesn't exist

An assistant suggested this to compute the average of a list:

```python
import statistics

def average(values):
    return statistics.average(values)
```

It looks reasonable, but it crashes. **Hallucinated APIs** — functions or
arguments that sound right but do not exist — are one of the most common failure
modes. Write a working `average(values)` that returns the mean of `values`.

<!-- grade_id: cell-ex04-bugA-sol, points: 0, grade: false, solution: true -->
```python
import statistics

def average(values):
    return statistics.mean(values)
```

<!-- grade_id: cell-ex04-bugA-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your average() is correct.
assert ads_hash(round(average([1, 2, 3, 4]) * 1000)) == '5a0b83e19c'
assert ads_hash(round(average([10, 20, 30]) * 1000)) == '876c9b1625'
```

### Bug B — plausible, but the logic is wrong

This function is supposed to return the **fraction** of values that are strictly
greater than `threshold`. It runs without error and even returns something that
looks like a fraction — but it is wrong.

```python
def proportion_above(values, threshold):
    count = 0
    for v in values:
        if v > threshold:
            count += 1
    return count / threshold
```

Find the mistake and write a correct `proportion_above(values, threshold)`.
(Check it by hand: of `[1, 2, 3, 4, 5]`, two values are above `3`, so the answer
should be `2/5 = 0.4`.)

<!-- grade_id: cell-ex04-bugB-sol, points: 0, grade: false, solution: true -->
```python
def proportion_above(values, threshold):
    count = 0
    for v in values:
        if v > threshold:
            count += 1
    return count / len(values)
```

<!-- grade_id: cell-ex04-bugB-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your proportion_above() is correct.
assert ads_hash(round(proportion_above([1, 2, 3, 4, 5], 3) * 1000)) == '26d228663f'
assert ads_hash(round(proportion_above([10, 10, 10, 10], 5) * 1000)) == '4051017584'
```

### Bug C — a silent type error

Suppose body-mass values were read from a file as **strings** (as they first are,
before you convert them). This function is meant to return the largest mass as a
number:

```python
def heaviest(masses_as_strings):
    return max(masses_as_strings)
```

Try it in your head on `['900', '1000', '3500']`. `max` on *strings* compares them
alphabetically, so it returns `'900'` — no error, just a wrong answer. This kind
of **silent** mistake is dangerous precisely because nothing complains. Write a
`heaviest(masses_as_strings)` that returns the largest value as a number.

<!-- grade_id: cell-ex04-bugC-sol, points: 0, grade: false, solution: true -->
```python
def heaviest(masses_as_strings):
    return max(float(m) for m in masses_as_strings)
```

<!-- grade_id: cell-ex04-bugC-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your heaviest() compares numerically.
# (float() is applied here so a correct int- or float-returning answer both pass.)
assert ads_hash(float(heaviest(['900', '1000', '3500']))) == '4a331d5936'
assert ads_hash(float(heaviest(['4200', '980', '5000', '760']))) == 'e0d61b8bd8'
```

## Part 4 — A plot, verified by its data

Ask the assistant to help you make a scatter plot of **culmen length** (x-axis)
against **body mass** (y-axis), using the `data` dictionary you built in Part 2.
Remember to leave out the missing-data rows (the `-999` sentinels), otherwise your
plot will have penguins that weigh almost nothing.

A plot is easy to trust because it *looks* like something. But you should verify
the data going *into* it. Build two lists first:

- `plot_x` — the culmen lengths of the valid rows,
- `plot_y` — the body masses of the valid rows,

then make the scatter plot from them (label your axes). The check cell verifies
`plot_x` and `plot_y`.

<!-- grade_id: cell-ex04-plot-sol, points: 0, grade: false, solution: true -->
```python
plot_x = []
plot_y = []
for culmen, mass in zip(data['culmen_length_mm'], data['body_mass_g (g)']):
    if mass != -999:
        plot_x.append(culmen)
        plot_y.append(mass)

plt.plot(plot_x, plot_y, '.')
plt.xlabel("Culmen length (mm)")
plt.ylabel("Body mass (g)");
```

<!-- grade_id: cell-ex04-plot-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, the data behind your plot is correct.
assert ads_hash(len(plot_x)) == '058d5d43bf'          # 334 valid rows
assert ads_hash(len(plot_y)) == '058d5d43bf'
assert ads_hash(round(sum(plot_y))) == '275807f1ea'   # sentinels excluded
```

## Part 5 — The assistant at the command line

Assistants help just as much with shell commands as with Python. Ask the
assistant for a **single shell command** that counts how many lines in
`antarctic_penguins.csv` belong to the `Adelie` species, and run it in the cell
below (in a notebook, a line starting with `!` runs a shell command). On macOS or
Linux the assistant will probably suggest something with `grep`; on Windows it
might suggest `findstr`. This cell is **not** graded — it is here so you can
compare the shell's answer with Python's.

```python
# Example (macOS / Linux). Replace with whatever the assistant gives you.
!grep -c Adelie antarctic_penguins.csv
# On Windows, try:  !findstr /R Adelie antarctic_penguins.csv | find /c /v ""
```

### Q1 — Now do it in Python and verify

Whatever the shell reported, verify it in Python. Using the `data` dictionary,
count how many penguins belong to the `Adelie` species and store the result in a
variable `n_adelie`. Does it match the shell's count?

<!-- grade_id: cell-ex04-cli-sol, points: 0, grade: false, solution: true -->
```python
n_adelie = 0
for species in data['species']:
    if species == "Adelie":
        n_adelie += 1
```

<!-- grade_id: cell-ex04-cli-test, points: 1, grade: true, solution: false -->
```python
# If this runs without error, your Adelie count is correct.
assert ads_hash(n_adelie) == '8e612bd1f5'
```

## Part 6 — Provenance and integrity (written answer, graded by a human)

Using an AI assistant well means being honest and reproducible about it. Answer
**both** of the following in the cell below:

1. Which model did you use (e.g. GPT-5.3, GLM-5.2, Gemma 4)? Paste the single
   prompt you found **most useful** during this exercise.
2. Describe **one** thing the assistant got wrong that you caught, and how you
   knew it was wrong.

<!-- grade_id: cell-ex04-reflect2, points: 1, grade: true, solution: true -->
*Replace this text with your answer.*

# Uploading the exercise

Upload this `.ipynb` file to ILIAS. **Do not rename the file.** Make sure you have
run every cell (top to bottom) so that your answers and any plots are saved in the
notebook before you upload it.

You do **not** need to upload `antarctic_penguins.csv`.