## The Jupyter notebook environment (formerly IPython notebooks)

Articles:

- https://www.dataquest.io/blog/jupyter-notebook-tutorial/

## Jupyter notebooks let you execute Python code in a browser.

They have advantages and disadvantages compared to other ways of running
Python code. We will start with them in this class to rapidly get programming.

In this course, we use [Sunbeam](https://github.com/strawlab/sunbeam) to
launch Jupyter Lab with Python already set up on your own computer.

```python
# this is a comment
2+2
```

Output:
```
4
```

```python
a = 4
```

The result of the last line is displayed in the notebook environment.

```python
a
```

Output:
```
4
```

```python
a = 5
```

```python
a
```

Output:
```
5
```

## Good practice

It is good practice to "Kernel -> restart and run all" prior to finishing work on a Jupyter notebook.

In this course, we are also using the `nbgrader` Jupyter extension. This adds a `Validate` button at the top of the window which re-runs the notebook.

<!-- grade_id: cell-2ff2ea3096b36c93, points: 1, grade: true, solution: false -->
```python
# This is an nbgrader automatic test.

assert a==5
```

If you do not see the "Validate" button above, `nbgrader` may not be
installed in your environment. If necessary, install it in a terminal with
`uv add nbgrader` (or ask an instructor for help).

## Jupyter Lab

```python

```