# Python programming basics

Here are some basic examples of concepts that we will learn in the class. We will also run these at http://pythontutor.com/visualize.html

## Built-in names

Python has several words which are "built in" and do special things. When we start, we will use a lot of these. As we get further along, we will start using our own names and names we *imported* more and more.

```python
print("hello world")
```

Output:
```
hello world
```

## Variables

Another building block of programming are **variables**. These are names given to hold values. Each variable has a type, such as a string (like "hello" -- note the quotation marks) or integer (like 42).

```python
x='hello world'
```

```python
print(x)
```

Output:
```
hello world
```

```python
print("x")
```

Output:
```
x
```

```python
x
```

Output:
```
'hello world'
```

```python
type(x)
```

Output:
```
str
```

```python
x=42
```

```python
print(x)
```

Output:
```
42
```

```python
print(42)
```

Output:
```
42
```

```python
x
```

Output:
```
42
```

```python
type(x)
```

Output:
```
int
```

```python
print(print)
```

Output:
```
<built-in function print>
```

```python
print
```

Output:
```
<function print(*args, sep=' ', end='\n', file=None, flush=False)>
```

```python
type(print)
```

Output:
```
builtin_function_or_method
```

```python
print(type)
```

Output:
```
<class 'type'>
```

```python
type
```

Output:
```
type
```

```python
type(type)
```

Output:
```
type
```

## Assignment

The process of setting a variable is called **assignment**. This is done with the equals (`=`) symbol.

```python
x = "my new string"
```

```python
x
```

Output:
```
'my new string'
```

```python
x = -1
```

```python
x
```

Output:
```
-1
```

```python
x = x + 1
```

```python
x
```

Output:
```
0
```

Let's look at these in the Python Tutor...

## Functions

A building block of programming are **functions**. These are pieces of code that are called (also known as executed) with input and return output. They may also have side effects.

Let's look at a function definition:

```python
def my_function_name(argument1, argument2):
    variable1 = argument1 + argument2
    print("hello from my function") # this is a side effect
    return variable1
```

Above, we created a function named `my_function_name`. This function takes two inputs, called arguments. There are several steps performed in the function. First, we add `argument1` to `argument2` and store the result in `variable1`. Then we print something. Then we return `variable1` as the output of our function.

The indentation (spaces) at the beginning of the lines in the function tell Python which lines of code are part of the function. Everything with the same level of indentation is a block.

Above, we called the built-in function `print()`.

The `print()` function takes any number of arguments as input, returns `None`, and as a side effect prints the arguments to the screen.

```python
def triple(x):
    return x * 3
```

```python
triple(4)
```

Output:
```
12
```

A function can have more than one argument.

```python
def weighted_sum(a, b):
    return 5 * a + 3 * b
```

```python
weighted_sum(3, 10)
```

Output:
```
45
```

A function can also call another function. Predict the result before running this cell.

```python
def triple_thrice(x):
    y = triple(x)
    z = triple(y)
    return triple(z)

triple_thrice(2)
```

Output:
```
54
```

Variables created inside a function live in the function's local scope. They do not replace variables with the same name outside the function.

```python
x = 43

def subtract_inside_function(x):
    x = x - 26
    return x

inside_value = subtract_inside_function(2)
print(inside_value)
print(x)
```

Output:
```
-24
43
```

A common beginner mistake is to print a value when the task asks a function to return a value.

```python
def prints_but_returns_none(x):
    print(x * 30)

result = prints_but_returns_none(10)
print(result)
```

Output:
```
300
None
```

## The predict → run → inspect → debug cycle

Programming is rarely about writing correct code on the first try. Instead,
it's an iterative loop, and it's worth naming explicitly because you'll repeat
it constantly for the rest of this course (and your career):

1. **Predict** what a piece of code will do before you run it — try this with
   the Python Tutor above, guessing the value of a variable before clicking
   "Next".
2. **Run** it.
3. **Inspect** what actually happened — the printed output, the variable
   values, or an error message.
4. **Debug** — reconcile any difference between what you predicted and what
   happened, then repeat.

The three cells below are a good place to practice: before running each one,
predict what will happen.

## Errors ("exceptions")

Python errors are called "exceptions". An important part of programming is figuring out why you got a particular error. Read the error message very carefully - it contains type of error, a description of the error, and a "traceback" that shows where the error came from.

```python
1/0
```

Output:
```
ZeroDivisionError                         Traceback (most recent call last)
Cell In[33], line 1
----> 1 1/0

ZeroDivisionError: division by zero
```

```python
x = variable_does_not_exist + 4
```

Output:
```
NameError                                 Traceback (most recent call last)
Cell In[34], line 1
----> 1 x = variable_does_not_exist + 4

NameError: name 'variable_does_not_exist' is not defined
```

```python
defx asdlkfj39:
    adsk..{
```

Output:
```
Cell In[35], line 1
    defx asdlkfj39:
         ^
SyntaxError: invalid syntax
```

## Doing the exercises

We use [nbgrader](https://nbgrader.readthedocs.io/en/stable/) for this course.
This way you can check most of your own work. Below is an example of how this
works. You can check the "Validate" button in the Jupyter Notebook to ensure
things are working OK.

When you have finished your exercise, upload the `.ipynb` file to ILIAS before
the deadline.

## Example Problem

Assign a value of 2 to the variable named `x`.

```python
x =2
```

<!-- grade_id: cell-a83e7ba665c8308e, points: 1, grade: true, solution: false -->
```python
assert x==2
```